Effects of blood meal source and seasonality on reproductive traits of Culex quinquefasciatus (Diptera: Culicidae)

Author response:

The following is the authors’ response to the previous reviews

We would like to thank you for your valuable comments and suggestions, which have greatly contributed to improving our manuscript.

We have carefully addressed all the reviewers' suggestions, and detailed responses for each Reviewer are provided at the end of this letter. In summary:

• The Introduction has been revised to provide a more focused discussion on results, toning down the speculative discussion on seasonal host shifts.

• The methodology section has been clarified, particularly the power analysis, which now includes a clearer explanation. The random effects in the models have been better described to ensure transparency.

• The Results section was reorganized to highlight the key findings more effectively.

• The Discussion has been restructured for clarity and conciseness, ensuring the interpretation of the results is clearer and better aligned with the study objectives.

• Minor edits throughout the manuscript were made to improve readability and accuracy.

We hope you find this revised version of the manuscript satisfactory.

Reviewer #1 (Public review):

Summary:

This study examines the role of host blood meal source, temperature, and photoperiod on the reproductive traits of Cx. quinquefasciatus, an important vector of numerous pathogens of medical importance. The host use pattern of Cx. quinquefasciatus is interesting in that it feeds on birds during spring and shifts to feeding on mammals towards fall. Various hypotheses have been proposed to explain the seasonal shift in host use in this species but have provided limited evidence. This study examines whether the shifting of host classes from birds to mammals towards autumn offers any reproductive advantages to Cx.

quinquefasciatus in terms of enhanced fecundity, fertility, and hatchability of the offspring. The authors found no evidence of this, suggesting that alternate mechanisms may drive the seasonal shift in host use in Cx. quinquefasciatus.

Strengths:

Host blood meal source, temperature, and photoperiod were all examined together.

Weaknesses:

The study was conducted in laboratory conditions with a local population of Cx. quinquefasciatus from Argentina. I'm not sure if there is any evidence for a seasonal shift in the host use pattern in Cx. quinquefasciatus populations from the southern latitudes.

Comments on the revision:

Overall, the manuscript is much improved. However, the introduction and parts of the discussion that talk about addressing the question of seasonal shift in host use pattern of Cx. quin are still way too strong and must be toned down. There is no strong evidence to show this host shift in Argentinian mosquito populations. Therefore, it is just misleading. I suggest removing all this and sticking to discussing only the effects of blood meal source and seasonality on the reproductive outcomes of Cx. quin.

Introduction and discussion have been modified, toned down and sticked to discuss the results as suggested.

Reviewer #1 (Recommendations for the authors):

Some more minor comments are mentioned below.

Line 51: Because 'of' this,

Changed as suggested.

Line 56: specialists 'or' generalists

Changed as suggested.

Line 56: primarily

Changed as suggested.

Line 98: Because 'of' this,

Changed as suggested.

Reviewer #2 (Public review):

Summary:

Conceptually, this study is interesting and is the first attempt to account for the potentially interactive effects of seasonality and blood source on mosquito fitness, which the authors frame as a possible explanation for previously observed hostswitching of Culex quinquefasciatus from birds to mammals in the fall. The authors hypothesize that if changes in fitness by blood source change between seasons, higher fitness on birds in the summer and on mammals in the autumn could drive observed host switching. To test this, the authors fed individuals from a colony of Cx. quinquefasciatus on chickens (bird model) and mice (mammal model) and subjected each of these two groups to two different environmental conditions reflecting the high and low temperatures and photoperiod experienced in summer and autumn in Córdoba, Argentina (aka seasonality). They measured fecundity, fertility, and hatchability over two gonotrophic cycles. The authors then used generalized linear mixed models to evaluate the impact of host species, seasonality, and gonotrophic cycle on fecundity, fertility, and hatchability. The authors were trying to test their hypothesis by determining whether there was an interactive effect of season and host species on mosquito fitness. This is an interesting hypothesis; if it had been supported, it would provide support for a new mechanism driving host switching. While the authors did report an interactive impact of seasonality and host species, the directionality of the effect was the opposite from that hypothesized. The authors have done a very good job of addressing many of the reviewer's concerns, especially by adding two additional replicates. Several minor concerns remain, especially regarding unclear statements in the discussion.

Strengths:

(1) Using a combination of laboratory feedings and incubators to simulate seasonal environmental conditions is a good, controlled way to assess the potentially interactive impact of host species and seasonality on the fitness of Culex quinquefasciatus in the lab.

(2) The driving hypothesis is an interesting and creative way to think about a potential driver of host switching observed in the field.

Weaknesses:

(1) The methods would be improved by some additional details. For example, clarifying the number of generations for which mosquitoes were maintained in colony (which was changed from 20 to several) and whether replicates were conducted at different time points.

Changed as suggested.

(2) The statistical analysis requires some additional explanation. For example, you suggest that the power analysis was conducted a priori, but this was not mentioned in your first two drafts, so I wonder if it was actually conducted after the first replicate. It would be helpful to include further detail, such as how the parameters were estimated. Also, it would be helpful to clarify why replicate was included as a random effect for fecundity and fertility but as a fixed effect for hatchability. This might explain why there were no significant differences for hatchability given that you were estimating for more parameters.

The power analysis was conducted a posteriori, as you correctly inferred. While I did not indicate that it was performed a priori, you are right in noting that this was not explicitly mentioned. As you suggested, the methodology for the power analysis has been revised to clarify any potential doubts.

Regarding the model for hatchability, a model without a random effect variable was used, as all attempts to fit models with random effects resulted in poor validation. These points have now been clarified and explained in the corresponding section.

(3) A number of statements in the discussion are not clear. For example, what do you mean by a mixed perspective in the first paragraph? Also, why is the expectation mentioned in the second paragraph different from the hypothesis you described in your introduction?

Changed as suggested.

(4) According to eLife policy, data must be made freely available (not just upon request).

Data and code will be publicly available. The corresponding section was modified.

Reviewer #2 (Recommendations for the authors):

Your manuscript is much improved by the inclusion of two additional replicates! The results are much more robust when we can see that the trends that you report are replicable across 3 iterations of the experiment. Congratulations on a greatly improved study and paper! I have several minor concerns and suggestions, listed below:

38-39: I think it is clearer to say "no statistically significant effect of season on hatchability of eggs" ... or specify if you are referring to blood or the interaction of blood and season. It isn't clear which treatment you are referring to here.

Changed as suggested.

54-57: This could be stated more succinctly. Instead of citing papers that deal with specific examples of patterns, I would suggest citing a review paper that defines these terms.

Changed as suggested.

83-84: What if another migratory bird is the preferred host in Argentina? I would state this more cautiously (e.g. "may not be applicable...").

Changed as suggested.

95-96: I don't understand what you mean by this. These hypotheses are specifically meant to understand mosquitoes that DO have a distinct seasonal phenology, so I'm not sure why this caveat is relevant. And naturally this hypothesis is host dependent, since it is based on specific host reproductive investments. I think that the strongest caveat to this hypothesis is simply that it hasn't been proven.

Changed as suggested.

97-115: This is a great paragraph! Very clear and compelling.

Thanks for your words!

118: Do you have an exact or estimated number of rafts collected?

Sorry, I have not the exact number of rafts, but it was at leas more than 20-30.

135: "over twenty" was changed to "several"; several would imply about 3 generations, so this is misleading. If the colony was actually maintained for over twenty generations, then you should keep that wording.

Changed as suggested.

163-164: Can you please clarify whether the replicates were conducted a separate time points?

Changed as suggested.

Note: the track changes did not capture all of the changes made; e.g. 163-164 should show as new text but does not.

You are absolutely right; when I uploaded the last version, I unfortunately deleted all tracked changes and cannot recover them. In this new version, I will ensure that all minimal changes are included as tracked changes.

186 - 189: the terms should be "fixed effect" and "random effect"

Changed as suggested.

191: Edit: linear

Changed as suggested.

194: why was replicate not included as a random effect here when it was above? Also, can you please clarify "interaction effects"? Which interactions did you include?

Changed as suggested. Explained above and in methodology. Hatchability models with random effect variable were poor fitted and validated. The interactions for hatchability were a four-way (season, blood source, cycle and replicate)

207-208: I'm not sure what you mean by "aimed to achieve"? Weren't you doing this after you conducted the experiments, so wouldn't this be determining the power of your model (post-hoc power analysis)? Also, I think you should provide the parameter estimates that were used (e.g. effect size - did you use the effect size you estimated across the 3 replicates?).

Changed as suggested.

214-215: this should be reworded to acknowledge that this is estimated for the given effect size; for example, something like "This sample size was sufficient to detect the observed effect with a statistical power of 0.8" or something along those lines (unless I am misunderstanding how you conducted this test).

Changed as suggested.

246. Abbreviate Culex

Changed as suggested.

253-255: This sentence isn't clear. What do you mean by mixed? Also, the season really seemed to mainly impact the fitness of mosquitoes fed on mouse blood and here the way it is phrased seems to indicate that season has an impact on the fitness of those fed with chicken blood.

Changed as suggested.

258-260: You stated your hypothesis as the relative fitness shifting between seasons, but this statement about the expectation is different from your hypothesis stated earlier. Please clarify.

You are right. Thank you for noting this. It was changed as suggested.

263-266: I also don't understand this sentence; what does the first half of the sentence have to do with the second?

Changed as suggested.

269-270: This doesn't align with your observation exactly; you say first AND second are generally most productive, but you observed a drop in the second. Please clarify this.

Changed as suggested.

280: I suggest removing "as same as other studies"; your caveats are distinct because your experimental design was unique

Changed as suggested.

287: you shouldn't be looking for a "desired" effect; I suggest removing this word

Changed as suggested.

288: It wasn't really a priori though, since you conducted it after your first replicate (unless you didn't use the results from the first replicate you reported in the original drafts?)

It was a posteriori. Changed as suggested.

290: Why is 290 written here?

It was a mistype. Deleted as suggested.

291-298: The meaning of this section of your paragraph is not clear.

Improve as suggested.

304-313: This list of 3 explanations are directed at different underlying questions. Explanations 1 and 2 are alternative explanations for why host switching occurs if not due to differences in fitness. This isn't really an explanation of your results so much as alternative explanations for a previously reported phenomenon. And the third is an explanation for why you may not have observed the expected effect. I suggest restructuring this to include the fact that Argentinian quinqs may not host switch as part of your previous list of caveats. Then you can include your two alternative explanations for host switching as a possible future direction (although I would say that it is really just one explanation because "vector biology" is too broad of a statement to be testable). Also, you haven't discussed possible explanations for your actual result, which showed that mosquito fitness decreased when feeding on mouse blood in autumn conditions and in the second gonotrophic, while those that fed on chicken did not experience these changes. Why might that be?

The discussion was restructured to include all these suggested changes. Additionally, it was also discussed some possible explanations of our results.

315-317: This statement is vague without a direct explanation of how this will provide insight. I suggest removing or providing an explanation of how this provides insight to transmission and forecasting.

Changed as suggested.

319-320: According to eLife policy, all data should be publicly available. From guidelines: "Media Policy FAQs Data Availability Purpose and General Principles To maintain high standards of research reproducibility, and to promote the reuse of new findings, eLife requires all data associated with an article to be made freely and widely available. These must be in the most useful formats and according to the relevant reporting standards, unless there are compelling legal or ethical reasons to restrict access. The provision of data should comply with FAIR principles (Findable, Accessible, Interoperable, Reusable). Specifically, authors must make all original data used to support the claims of the paper, or that is required to reproduce them, available in the manuscript text, tables, figures or supplementary materials, or at a trusted digital repository (the latter is recommended). This must include all variables, treatment conditions, and observations described in the manuscript. The authors must also provide a full account of the materials and procedures used to collect, pre-process, clean, generate and analyze the data that would enable it to be independently reproduced by other researchers."

- so you need to make your data available online; I also understand the last sentence to indicate that code should be made available.

Data and code will be publicly available.

Table 1: it is notable that in replicate 2, the autumn:mouse:gonotrophic cycle II fecundity and fertility are actually higher than in the summer, which is the opposite of reps 1 and 3 and the overall effect you reported from the model. This might be worth mentioning in the discussion.

Mentioned in the discussion as suggested.

Tables 1 and 2: shouldn't this just be 8 treatments? You included replicate as a random effect, so it isn't really a separate set of treatments.

This table reflects the output of the whole experiment, that is why it is present the 24 expetiments.

Figure 3: Can you please clarify if this is showing raw data?

Changed as suggested.

Note: grammatical copy editing would be beneficial throughout

Grammar was improved as suggested.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Host selection by mosquitoes is a keystone to understanding viral circulation and predicting future infection outbreaks. Culex mosquitoes frequently feed on birds during spring and early summer, shifting into mammals towards late summer and autumn. This host switch may be due to changes in mosquito fitness. The aim of this study was to assess if the interaction effect of blood meal source and seasonality may influence reproductive traits of Culex quinquefasciatus mosquitoes. For this purpose, Cx. quinquefasciatus mosquitoes were reared in simulated summer and autumn conditions and fed on two different hosts, chickens and mice, in a factorial design. Fecundity, fertility, and hatchability during two consecutive gonotrophic cycles were estimated. We found greater fecundity and fertility for mosquitoes fed upon birds than mammals. Fecundity and fertility did not vary between seasons for chicken-fed mosquitoes, whereas in autumn they decreased for mouse-fed mosquitoes. These traits decreased in the second gonotrophic cycle for mouse-fed mosquitoes, whereas they did not vary between cycles for chicken-fed mosquitoes. There was no statistically significant variation of hatchability among treatments. These results indicate a statistically significant interaction effect of blood meal source and seasonality on fecundity and fertility. However, the pattern was opposite in relation to our hypothesis, suggesting that further studies are needed to confirm and expand our knowledge about mosquito biology and its relationship with seasonal host use shifting.

Revised version: Host selection by mosquitoes is a keystone to understanding viral circulation and predicting future infection outbreaks. Culex mosquitoes frequently feed on birds during spring and early summer, shifting into mammals towards late summer and autumn. This host switch may be due to changes in mosquito fitness. The aim of this study was to assess if the interaction effect of blood meal source and seasonality may influence reproductive traits of Culex quinquefasciatus mosquitoes. For this purpose, Cx. quinquefasciatus mosquitoes were reared in simulated summer and autumn conditions and fed on two different hosts, chickens and mice, in a factorial design. Fecundity, fertility, and hatchability during two consecutive gonotrophic cycles were estimated. We found greater fecundity and fertility for mosquitoes fed upon birds than mammals. Fecundity and fertility did not vary between seasons for chicken-fed mosquitoes, whereas in autumn they decreased for mouse-fed mosquitoes. These traits decreased in the second gonotrophic cycle for mouse-fed mosquitoes, whereas they did not vary between cycles for chicken-fed mosquitoes. There was no statistically significant effect of blood meal source, seasonality or their interaction on hatchability, hence this variable was similar among treatments. Overall, these results indicate a statistically significant interaction effect of blood meal source and seasonality on fecundity and fertility. However, the pattern was opposite in relation to our hypothesis, suggesting that further studies are needed to confirm and expand our knowledge about mosquito biology and its relationship with seasonal host use shifting.

Introduction

Earlier version: The southern house mosquito, Culex quinquefasciatus Say (Diptera: Culicidae), is a worldwide distributed vector of numerous pathogens, including Wuchereria bancrofti Cobbold (Kasili et al., 2009), Dirofilaria immitis Leidy (Labarthe et al., 1998) and avian malaria (Farajollahi et al., 2011). Besides, this mosquito is also one of the main species responsible for the transmission of the arboviruses West Nile (WNV) and St. Louis encephalitis (SLEV) viruses (Farajollahi et al., 2011). Because this, this species is of critical concern due to its impact on both public and veterinary health.

Revised version: The southern house mosquito, Culex quinquefasciatus Say (Diptera: Culicidae), is a worldwide distributed vector of numerous pathogens, including Wuchereria bancrofti Cobbold (Kasili et al., 2009), Dirofilaria immitis Leidy (Labarthe et al., 1998) and avian malaria (Farajollahi et al., 2011). Besides, this mosquito is also one of the main species responsible for the transmission of the arboviruses West Nile (WNV) and St. Louis encephalitis (SLEV) viruses (Farajollahi et al., 2011). Because of this, this species is of critical concern due to its impact on both public and veterinary health.

Earlier version: Since arboviruses are transmitted by the bite of infected mosquitoes between vertebrate hosts, activity of these pathogens is determined, at least in part, by the blood feeding habits of these vectors. According to its host feeding pattern range, Culex mosquitoes can be regarded as specialists or generalists. Specialist species include mosquitoes feeding primarily on mammals (Muturi et al., 2008), birds (Jansen et al., 2009), reptiles and amphibians (Janssen et al., 2015). Generalists are mosquitoes fed on the most abundant available hosts (Molaei et al., 2007). Additionally, some species could experience a seasonal shift in its feeding habits, from avian to mammal hosts (Edman & Taylor, 1968; Hancock & Camp, 2022), being of greater concern from an epidemiological perspective given that they can act as bridge vectors in the transmission of some arboviruses, such as SLEV and WNV (Kilpatrick et al., 2006).

Revised version: Since arboviruses are transmitted by the bite of infected mosquitoes between vertebrate hosts, activity of these pathogens is determined, at least in part, by the blood feeding habits of these vectors. According to its host feeding pattern range, Culex mosquitoes can be regarded as specialists or generalists. Specialist species include mosquitoes feeding primarily on mammals, birds, or reptiles and amphibians. Generalists are mosquitoes fed on the most abundant available hosts (Fikrig and Harrington, 2021). Additionally, some species could experience a seasonal shift in their feeding habits, from avian to mammal hosts (Edman and Taylor, 1968; Hancock and Camp, 2022), being of greater concern from an epidemiological perspective given that they can act as bridge vectors in the transmission of some arboviruses, such as SLEV and WNV (Kilpatrick et al., 2006).

Earlier version: West Nile virus (WNV) and St. Louis encephalitis virus (SLEV) are zoonotic pathogens maintained in nature through an enzootic network involving transmission among viremic birds and ornithophagic mosquitoes (Kilpatrick et al., 2006). However, during late summer and early autumn, these viruses increase their activity in humans and other mammals, which has led to repeated epidemics in the US (Kilpatrick et al., 2006) and South America (Spinsanti et al., 2008). Likewise, several Culex species, including Cx. quinquefasciatus, have exhibited a seasonal shift in their feeding behaviour, feeding primarily on birds in spring and early summer and increasing the proportion of mammal blood meals in late summer and autumn (Thiemann et al., 2011). This seasonal shift suggests that viral activity might spill over from birds to mammals as a result of the mosquito host switch (Edman & Taylor, 1968; Kilpatrick et al., 2006).

Revised version: West Nile virus (WNV) and St. Louis encephalitis virus (SLEV) are zoonotic pathogens maintained in nature through an enzootic network involving transmission among viremic birds and ornithophagic mosquitoes (Kilpatrick et al., 2006). However, during late summer and early autumn, these viruses increase their activity in humans and other mammals, which has led to repeated epidemics in the US (Kilpatrick et al., 2006) and South America (Spinsanti et al., 2008). Likewise, several Culex species, including Cx. quinquefasciatus, have exhibited a seasonal shift in their feeding behaviour, feeding primarily on birds in spring and early summer and increasing the proportion of mammal blood meals in late summer and autumn (Thiemann et al., 2011). This seasonal shift suggests that viral activity might spill over from birds to mammals as a result of the mosquito host switch (Edman and Taylor, 1968; Kilpatrick et al., 2006).

Earlier version: Two main contrasting hypotheses have been proposed to explain why primarily ornithophagic mosquitoes increase their feeding on mammals later in the season. The migration hypothesis, proposed by Kilpatrick et al. (2006) for Cx. pipiens, suggests that this shift in feeding pattern is related to the abundance of the preferred avian host, the American robin (Turdus migratorius L.). This hypothesis is based on their findings that as robin densities decrease, the proportion of mammalian blood meals in Cx. pipiens populations increase. However, this hypothesis has several limitations: (i) multiple studies have shown that American robins are not always the preferred host for all Cx. pipiens populations (Patrican et al., 2007; Apperson et al., 2002); (ii) in some areas (e.g., Tennessee) with documented seasonal host shifts, American robins are present year-round (Savage et al., 2007); and (iii) this explanation applies only to mosquitoes whose primary host is the American robin, a species found only in the Northern Hemisphere, or another migratory bird. Therefore, it is not applicable to southern regions, such as Argentina, where SLEV seasonality is documented (Spinsanti et al., 2008) but American robins are absent. Because this, it seems that the hypothesis proposed by Kilpatrick et al. (2006) might be habitat dependent (Caranci, 2010). On the other hand, the defensive behaviour hypothesis proposed by Burkett-Cadena et al. (2011), suggests that host breeding cycles drive the host shift of mosquitoes. According to this hypothesis, in hosts with parental care, during reproductive seasons there is a greater investment of energy in assuring the offspring survivorship, consequently producing an increase in susceptibility of being bitten by mosquitoes as a result of a decrease in defensive behaviours. This event leads to the detection of peaks of host use during periods of reproductive investment, in summer for birds and autumn for mammals (Burkett-Cadena et al., 2011). While this hypothesis was proposed as an alternative to address the limitations of the migration hypothesis, it also has drawbacks. It is based on a limited range of host species and assumes a distinct seasonal phenology, which is not present for all mosquito hosts. Because this, this hypothesis seems to be host dependent.

Revised version: Two main contrasting hypotheses have been proposed to explain why primarily ornithophagic mosquitoes increase their feeding on mammals later in the season. The migration hypothesis, proposed by Kilpatrick et al. (2006) for Cx. pipiens, suggests that this shift in feeding pattern is related to the abundance of the preferred avian host, the American robin (Turdus migratorius L.). This hypothesis is based on their findings that as robin densities decrease, the proportion of mammalian blood meals in Cx. pipiens populations increase. However, this hypothesis has several limitations: (i) multiple studies have shown that American robins are not always the preferred host for all Cx. pipiens populations (Apperson et al., 2002; Patrican et al., 2007); (ii) in some areas (e.g., Tennessee) with documented seasonal host shifts, American robins are present year-round (Savage et al., 2007); and (iii) this explanation applies only to mosquitoes whose primary host is the American robin, a species found only in the Northern Hemisphere, or another migratory bird. Therefore, it might not be applicable to southern regions, such as Argentina, where WNV seasonality is documented (Spinsanti et al., 2008) but American robin is absent. Because of this, it seems that the hypothesis proposed by Kilpatrick et al. (2006) might be habitat dependent (Caranci, 2010). On the other hand, the defensive behaviour hypothesis proposed by Burkett-Cadena et al. (2011), suggests that host breeding cycles drive the host shift of mosquitoes. According to this hypothesis, in hosts with parental care, during reproductive seasons there is a greater investment of energy in assuring the offspring survivorship, consequently producing an increase in susceptibility of being bitten by mosquitoes as a result of a decrease in defensive behaviours. This event leads to the detection of peaks of host use during periods of reproductive investment, in summer for birds and autumn for mammals (Burkett-Cadena et al., 2011). While this hypothesis was proposed as an alternative to address the limitations of the migration hypothesis, it also has drawbacks. It is based on a limited range of host species, assumes a distinct seasonal phenology—which is not present for all mosquito hosts—and remains untested, leaving its validity uncertain.

Earlier version: Explanations for the seasonal shift in mosquito host feeding focused on vector biology are largely lacking in the literature. Numerous biological factors, such as stress, metabolic rate, and blood meal source, along with environmental variables like temperature and photoperiod, influence mosquito physiology and can affect reproductive traits such as fecundity, development rate, and survivorship. These factors can give rise to new nutritional requirements that may ultimately lead to seasonal variations in host selection by mosquitoes (Ciota et al., 2014; Costanzo et al., 2015; Gervasi et al., 2016; Yan et al., 2017, 2018). While several studies have examined the effects of these biological and environmental variables on mosquito reproduction, the complex nature of mosquito biology suggests that multiple interactions among these variables may produce novel responses that are not apparent when considered individually. The aim of this study was to assess whether there is an interaction between the source of the host blood meal and seasonality (defined by temperature and photoperiod) on three reproductive traits of Culex quinquefasciatus mosquitoes: fecundity, fertility, and hatchability. We hypothesize that the interaction between these two variables influence the reproductive outcomes, potentially leading to a seasonal shift in host selection, driven by a reproductive advantage. Given the reported seasonal changes in host use by Cx. quinquefasciatus, in autumn we expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes after feeding on a mammal host compared to an avian host, and the opposite trend in summer.

Revised version: Explanations for the seasonal shift in mosquito host feeding focused on vector biology are largely lacking in the literature. Numerous biological factors, such as stress, metabolic rate, and blood meal source, along with environmental variables like temperature and photoperiod, influence mosquito physiology and can affect reproductive traits such as fecundity, development rate, and survivorship. These factors can give rise to new nutritional requirements that may ultimately lead to seasonal variations in host selection by mosquitoes (Ciota et al., 2014; Costanzo et al., 2015; Gervasi et al., 2016; Yan et al., 2018, 2017). While several studies have examined the effects of these biological and environmental variables on mosquito reproduction, the complex nature of mosquito biology suggests that multiple interactions among these variables may produce novel responses that are not apparent when considered individually. The aim of this study was to assess whether there is an interaction effect between the source of the host blood meal and seasonality (defined by temperature and photoperiod) on three reproductive traits of Culex quinquefasciatus mosquitoes: fecundity, fertility, and hatchability. We hypothesize that the interaction between these two variables influences the reproductive outcomes, potentially leading to a seasonal shift in host selection, driven by a reproductive advantage. Given the reported seasonal changes in host use by Cx. quinquefasciatus, in autumn we expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes after feeding on a mammal host compared to an avian host, and the opposite trend in summer.

Materials and methods

Establishment and maintenance of mosquitoes

Earlier version: Egg rafts of Culex quinquefasciatus were collected from a drainage ditch at Universidad Nacional de Córdoba Campus, Córdoba city, in February 2021. Each raft was individually maintained in plastic containers with one liter of distilled water. The hatched larvae were fed with 100 mg of liver powder three times per week until pupation. Pupae were then transferred to plastic emerging cages (21 cm x 12 cm) covered with a tulle-like fabric, containing distilled water but no food. Adults emerging from each raft were identified morphologically (Darsie, 1985) and molecularly (Smith & Fonseca, 2004) to ensure they corresponded to Cx. quinquefasciatus, since Cx. pipiens and its hybrids coexist sympatrically in Córdoba city (Branda et al., 2021). All adults were reared in a cardboard cage of 22.5 liters (28 cm x 36.5 cm) and provided ad libitum with a 10% sugar solution soaked in cotton pads placed on plastic cups. For long-term maintenance of the colony, 24-hour-starved mosquitoes were offered a blood meal from a restrained chicken twice a month. Four days after feeding, a plastic container with distilled water was placed inside the cage to allow engorged females to lay egg rafts. Batches of egg were collected and transferred to plastic containers (30 cm x 25 cm x 7 cm) in a proportion of three rafts per container, filled with three liters of distilled water. The hatched larvae were also fed with liver powder at the same proportion described above. Pupae were transferred to emerging cages, and adult mosquitoes were placed in the final cardboard cage. The colony has been maintained for several generations in the Insectary of the Instituto de Virología “Dr. J.M. Vanella” (InViV). Room conditions were controlled at 28 ± 1 °C, with a 12L:12D photoperiod and 70% relative humidity.

Revised version: Egg rafts of Culex quinquefasciatus were collected from a drainage ditch at Universidad Nacional de Córdoba Campus, Córdoba city, in February 2021. Each raft was individually maintained in plastic containers with one liter of distilled water. The hatched larvae were fed with 100 mg of liver powder three times per week until pupation. Pupae were then transferred to plastic emerging cages (21 cm x 12 cm) covered with a tulle-like fabric, containing distilled water but no food. Adults emerging from each raft were identified morphologically (Darsie, 1985) and molecularly (Smith and Fonseca, 2004) to ensure they corresponded to Cx. quinquefasciatus, since Cx. pipiens and its hybrids coexist sympatrically in Córdoba city (Branda et al., 2021). All adults were reared in a cardboard cage of 22.5 liters (28 cm x 36.5 cm) and provided ad libitum with a 10% sugar solution soaked in cotton pads placed on plastic cups. For long-term maintenance of the colony, 24-hour-starved mosquitoes were offered a blood meal from a restrained chicken twice a month. Four days after feeding, a plastic container with distilled water was placed inside the cage to allow engorged females to lay egg rafts. Batches of egg were collected and transferred to plastic containers (30 cm x 25 cm x 7 cm) in a proportion of three rafts per container, filled with three liters of distilled water. The hatched larvae were also fed with liver powder at the same proportion described above. Pupae were transferred to emerging cages, and adult mosquitoes were placed in the final cardboard cage. The colony has been maintained for over twenty generations in the Insectary of the Instituto de Virología “Dr. J.M. Vanella” (InViV). Room conditions were controlled at 28 ± 1 °C, with a 12L:12D photoperiod and 70% relative humidity.

Experimental design

Removed in the revised version: Blood source

Earlier version: For experimental trials, avian hosts (live chicks of the species Gallus gallus) and mammalian hosts (live mice of the species Mus musculus, strain C57BL/6) were used to evaluate the effect of blood meal source on fecundity, fertility, and hatchability. Chicks were generously donated by the Bartolucci poultry farm (Córdoba, Argentina), while mice were commercially obtained from the Instituto de Investigación Médica Mercedes y Martín Ferreyra (CONICET - Universidad Nacional de Córdoba). In each trial, 24-hour-starved adult female mosquitoes were provided with a blood meal from restrained chicks or mice. Vertebrate hosts were offered to mosquitoes one hour before the lights were turned off and were kept for 3 hours during 2 consecutive gonotrophic cycles.

Revised version: Blood source. For experimental trials, avian hosts (live chicks of the species Gallus gallus) and mammalian hosts (live mice of the species Mus musculus, strain C57BL/6) were used to evaluate the effect of blood meal source on fecundity, fertility, and hatchability. Chicks were generously donated by the Bartolucci poultry farm (Córdoba, Argentina), while mice were commercially obtained from the Instituto de Investigación Médica Mercedes y Martín Ferreyra (CONICET - Universidad Nacional de Córdoba). In each trial, 24-hour-starved adult female mosquitoes were provided with a blood meal from restrained chicks or mice. Vertebrate hosts were offered to mosquitoes one hour before the lights were turned off and were kept for 3 hours during 2 consecutive gonotrophic cycles.

The experimental use of animals (mosquitoes and vertebrates) was approved by the ethical committee at Facultad de Ciencias Médicas, Universidad Nacional de Córdoba (FCM-UNC) in compliance with the legislation regarding the use of animals for experimental and other scientific purposes (accession code: CE-2022-00518476-UNC-SCT#FCM).

Seasonality

To assess the effect of seasonality (photoperiod + temperature) on fecundity, fertility, and hatchability, a “typical” summer and autumn day from Córdoba city was simulated in an incubator, where mosquitoes were housed throughout the entire experiment. Summer conditions were as follows: T°min = 22°C, T°max = 28°C, photoperiod = 14L:10D. Autumn conditions were characterized by: T°min = 16°C, T°max = 22°C, and photoperiod = 10L:14D. The humidity levels were maintained at 60-70% for both conditions. Data for simulated conditions were obtained from Climate Data website (https://es.climate-data.org/).

Feeding trial design

Earlier version: The interaction effect of blood source (chicken and mouse) and seasonality (autumn and summer) was evaluated during two consecutive gonotrophic cycles (I and II). This combination produced eight colonies that were established using egg rafts collected from the maintenance colony: mouse:autumn-I, mouse:autumn-II, chicken:autumn-I, chicken:autumn-II, mouse:summer-I, mouse:summer-II, chicken:summer-I, chicken:summer-II. This set of 8 colonies was repeated in triplicate (1-3), generating a total of 24 treatments.

Revised version: The interaction effect of blood source (chicken or mouse) and seasonality (autumn or summer) was evaluated during two consecutive gonotrophic cycles (I and II). This combination produced eight colonies that were established using egg rafts collected from the maintenance colony: mouse:autumn-I, mouse:autumn-II, chicken:autumn-I, chicken:autumn-II, mouse:summer-I, mouse:summer-II, chicken:summer-I, chicken:summer-II. This set of eight colonies was repeated in triplicate (1-3) at different time points, generating a total of 24 treatments.

The experimental colonies were maintained under controlled conditions to ensure that all adults were of the same size (Kauffman et al., 2017). Adult mosquitoes were placed in cardboard cages of 8.3 liters (21 cm x 24 cm) and were provided with ad libitum access to a 10% sugar solution, which was soaked in cotton pads placed on plastic cups.

Earlier version: Feeding trials were conducted at two time points: 5 days post-emergence (first cycle) and 14 days post-emergence (second cycle). Following each blood meal, female mosquitoes were anesthetized using CO2 and classified as fully engorged (VI Sella’s stage), partially fed (I-V Sella’s stage), or unfed (I Sella’s stage), following the classification by Silva Santos et al. (2019). Fully engorged females were counted and separated into a separate cage to complete their gonotrophic cycle. Four days after feeding, a cup containing distilled water was placed inside the cage to facilitate oviposition. After each oviposition cycle, egg rafts were collected, counted, and transferred to a 12-well plate containing 4 mL of distilled water. They were then photographed to subsequently determine the number of eggs per raft. Rafts were maintained in each well until they hatched into L1 larvae, at which point they were preserved with 2 mL of 96% ethanol, and the number of L1 larvae per raft was counted.

Revised version: Feeding trials were conducted at two time points: 5 days post-emergence (first cycle) and 14 days post-emergence (second cycle). Following each blood meal, female mosquitoes were anesthetized using CO2 and classified as fully engorged (VI Sella’s stage), partially fed (I-V Sella’s stage), or unfed (I Sella’s stage), following the classification by Santos et al. (2019). Fully engorged females were counted and separated into a separate cage to complete their gonotrophic cycle. Four days after feeding, a cup containing distilled water was placed inside the cage to facilitate oviposition. After each oviposition cycle, egg rafts were collected, counted, and transferred to a 12-well plate containing 4 mL of distilled water. They were then photographed to subsequently determine the number of eggs per raft. Rafts were maintained in each well until they hatched into L1 larvae, at which point they were preserved with 2 mL of 96% ethanol, and the number of L1 larvae per raft was counted.

Statistical analysis

The reproductive outputs measured for each experimental colony were fecundity, fertility, and hatchability. All analyses were performed using R Studio statistic software, v.4.2.1 (R Core Team, 2022).

Earlier version: Fecundity was defined as the number of eggs per raft and fertility as the number of L1 larvae hatched per raft. To evaluate the effect of blood meal source and seasonality on these two variables, a Generalized Linear Mixed Model (GLMM; lme4 package) with negative binomial error distribution and logarithmic link function was adjusted (Bates et al., 2015). Fecundity and fertility served as the response variables, while the fixed explanatory variables included blood source, seasonality, and gonotrophic cycle, all in a three-way interaction. Treatments (n = 24) were set as a random intercept.

Revised version: Fecundity was defined as the number of eggs per raft and fertility as the number of L1 larvae hatched per raft. To evaluate the effect of blood meal source and seasonality on these two variables, a Generalized Linear Mixed Model (GLMM; lme4 package) with negative binomial error distribution and logarithmic link function was adjusted (Bates et al., 2015). Fecundity and fertility served as response variables, while the fixed effect variables included blood source, seasonality, and gonotrophic cycle, all in a three-way interaction. Treatment (n = 24) was set as a random effect variable.

Earlier version: Hatchability (= hatching rate) was defined as the ratio of the number of larvae to the number of eggs per raft. A Generalized Linear Model (GLM; MASS package) was applied using a quasipoisson error distribution and a logarithmic link function (Venables et al., 2002). The response variable was the number of larvae, while the explanatory variables included blood source, seasonality, gonotrophic cycle and replicate as interaction effects. Additionally, the logarithm of the number of eggs was incorporated as an offset in the model.

Revised version: Hatchability (= hatching rate) was defined as the ratio of the number of larvae to the number of eggs per raft. While the original intent was to apply a GLMM to all three variables, the validation of GLMMs for hatchability across several distributions revealed poor fit. Consequently, a Generalized Linear Model (GLM; MASS package) with a quasipoisson error distribution and logarithmic link function was used instead (Venables and Ripley, 2002). In this model, the response variable was the number of larvae, while the explanatory variables included blood source, seasonality, gonotrophic cycle, and replicate as a four-way interaction. Additionally, the logarithm of the number of eggs was incorporated as an offset in the model.

Earlier version: Multiple comparisons among treatments were conducted using Tukey’s honestly significant difference test (HSD Tukey), incorporating the Kramer (1956) correction for unbalanced data, also known as the Tukey-Kramer method. This approach allows setting the family-wise error rate to 0.05. These comparisons were performed using the emmeans package (Lenth, 2024).

Revised version: Multiple comparisons among treatments were conducted using the Tukey’s honestly significant difference test (HSD Tukey), incorporating the Kramer (1956) correction for unbalanced data, also known as the Tukey-Kramer method. This approach allows to set the family-wise error rate to 0.05. These comparisons were performed using the emmeans package (Lenth, 2017).

Earlier version: Model validation was assessed using a simulation-based approach to generate randomized quantile residuals to test goodness of fit (Q-Q plot), homoscedasticity, and outliers and influential points (Cook’s distance and Leverage plots). These analyses were conducted with the DHARMa package (Hartig, 2022).

Revised version: Model validation was assessed using a simulation-based approach to generate randomized quantile residuals to test goodness of fit (Q-Q plot), homoscedasticity, and outliers and influential points (Cook’s distance and Leverage plots). These analyses were conducted with the DHARMa package (Hartig, 2024).

Added in the revised version: Additionally, post-hoc statistical power was evaluated using the mixedpower package, a tool specifically designed for simulating statistical power in Generalized Linear Mixed Models (GLMMs). The analysis proceeded in three key steps: (1) simulation of new datasets based on the parameters of an original fitted model (the fecundity and fertility models), (2) refitting the model to each dataset and testing for statistical significance, and (3) calculating statistical power as the proportion of simulated datasets where the effect of interest reached significance. Fixed effects were simulated for sample sizes ranging from 30 to 150, with 1,000 datasets generated at each sample size to ensure robust estimation. The variable treatment was designated as a random effect, capturing variability associated with this factor. For further details on the methodology, including considerations for simulation accuracy and model structure, refer to Kumle et al. (2021).

Removed in the revised version: Additionally, the statistical power of the models was assessed using the mixedpower package (Kumle et al., 2021), specifically designed to evaluate the power of Generalized Linear Mixed Models (GLMMs). The analysis involved simulating 1,000 iterations with a sample size ranging from 30 to 150 observations per group. The power analysis aimed to achieve a power threshold of 0.8, indicating an 80% chance of detecting true effects if they exist.

Results

Earlier version: A total of 1,162 egg rafts were obtained for analysis, distributed across the eight treatments as follows: chicken:autumn-I: 260, chicken:autumn-II: 174, chicken:summer-I: 162, chicken:summer-II: 88, mouse:autumn-I: 146, mouse:autumn-II: 96, mouse:summer-I: 167, and mouse:summer-II: 69. The full count of egg rafts for each replicate (24 treatments), along with the means for fecundity, fertility, and hatchability, are summarized in Table 1. This sample size fulfilled the statistical power of 0.8 (Figure SM1).

Revised version: A total of 1,162 egg rafts were obtained for analysis, distributed across the eight treatments as follows: chicken:autumn-I: 260, chicken:autumn-II: 174, chicken:summer-I: 162, chicken:summer-II: 88, mouse:autumn-I: 146, mouse:autumn-II: 96, mouse:summer-I: 167, and mouse:summer-II: 69. The full count of egg rafts for each replicate (24 treatments), along with the means for fecundity, fertility, and hatchability, are summarized in Table 1. This sample size was sufficient to detect the observed effect with a statistical power of 0.8, indicating an 80% chance of detecting a true effect if it exists (Figure SM1).

In both models—fecundity and fertility—the interaction between blood source and seasonality was statistically significant (fecundity: LRT X² = 5.69, p < 0.05; fertility: LRT X² = 4.37, p < 0.05) (Tables 2,3; Figures 1,2A). Mosquitoes that fed on chicken blood had the highest fecundity (summer: 136.4 eggs/raft, autumn: 141.5 eggs/raft) and fertility (summer: 123.9 larvae/raft, autumn: 126 larvae/raft) in both seasons. In summer, fecundity and fertility were 7% and 11% higher, respectively, in chicken-fed mosquitoes compared to mouse-fed mosquitoes. In autumn, these differences increased to 46% for fecundity and 52% for fertility in favor of chicken-fed mosquitoes. No significant seasonal variation was observed in fecundity and fertility for chicken-fed mosquitoes, but both measures were statistically significant lower by 24% and 25%, respectively, in autumn compared to summer for mouse-fed mosquitoes.

Earlier version: Interaction plot showing the effect of blood meal source (chicken or mouse) and seasonality (autumn or summer) on fecundity (eggs/raft). Points represent the predicted marginal mean values and the bars corresponding to their confidence intervals. Means sharing the same letters are not statistically different (p > 0.05).

Revised version: Interaction plot showing the effect of blood meal source (chicken or mouse) and seasonality (autumn or summer) on fecundity (eggs/raft). Points represent the predicted marginal mean values and the bars its corresponding confidence intervals. Means sharing same letters are not statistically different (p > 0.05).

Earlier version: Interaction plots showing the effect of blood meal source (chicken and mouse) and seasonality (autumn and summer) (A) or blood meal source and gonotrophic cycle (first or second) (B) on fertility (eggs/raft) of Culex quinquefasciatus. Points represent the predicted marginal mean values and the bars corresponding to their confidence intervals. Means sharing the same letters are not statistically different (p > 0.05).

Revised version: Interaction plots showing the effect of blood meal source (chicken or mouse) and seasonality (autumn or summer) (A) or blood meal source and gonotrophic cycle (first or second) (B) on fertility (eggs/raft) of Culex quinquefasciatus. Points represent the predicted marginal mean values and the bars its corresponding confidence intervals. Means sharing same letters are not statistically different (p > 0.05).

In the fertility model, there was also a significant interaction between blood source and gonotrophic cycle (LRT X² = 3.98, p < 0.05), indicating that the impact of blood type on larvae production differed between the first and second cycles (Table 3, Figure 2B). Chicken-fed mosquitoes showed no difference in fertility between cycles, while those fed on mouse blood had 25% lower fertility in the second cycle compared to the first. During the second cycle, fertility was 50% higher in chicken-fed mosquitoes compared to mouse-fed mosquitoes.

Figure 3 presents the predicted mean values of fecundity and fertility for the 24 treatments, with no differences observed across the three replicates. The variance of random intercepts among treatments was approximately 0.006 for fecundity and 0.007 for fertility, accounting for 6% and 4% of the total variance in the respective models.

Earlier version: Dotplots showing the variation of predicted fecundity (eggs/raft) and fertility (larvae/raft) means along the 24 treatments, combination of blood source (chicken or mouse), seasonality (autumn or summer), gonotrophic cycle (first or second) and replicate (1 to 3).

Revised version: Dotplots showing the variation of predicted fecundity (eggs/raft) and fertility (larvae/raft) means along the 24 treatments, with combination of blood source (chicken or mouse), seasonality (autumn or summer), gonotrophic cycle (first or second) and replicate (1 to 3).

Earlier version: There was no statistical difference among treatments and replicates in terms of hatchability, given that no statistical effects of blood source, seasonality or gonotrophic cycle was observed (Table 4). The mean hatchability for all treatments was 0.89, with a range between 0.84 for the mouse:autumn-II (replicate 1) colony and 0.96 for mouse:autumn-I (replicate 1) (Table 1, Figure 4).

Revised version: There was no statistical difference among treatments and replicates in terms of hatchability, given that no statistical effects of blood source, seasonality or gonotrophic cycle were observed (Table 4). The mean hatchability for all treatments was 0.89, with a range between 0.84 for the mouse:autumn-II (replicate 1) colony and 0.96 for mouse:autumn-I (replicate 1) (Table 1, Figure 4).

Earlier version: Boxplots showing the mean hatchability for the 24 treatments, combining replicate (1-3), blood meal source (chicken and mouse), seasonality (summer and autumn) and gonotrophic cycle (first and second), of Culex quinquefasciatus.

Revised version: Boxplots showing predicted mean hatchability for the 24 treatments, combining replicate (1-3), blood meal source (chicken or mouse), seasonality (summer or autumn) and gonotrophic cycle (first or second), of Culex quinquefasciatus.

Discussion

Added in the revised version: Understanding the factors influencing mosquito host selection is crucial for predicting arbovirus transmission dynamics. Seasonal shift in host utilization have been widely attributed to changes in host behaviour (Burkett-Cadena et al., 2011; Kilpatrick et al., 2006), yet the role of mosquito biology in this process remains largely unexplored. Blood meal source and seasonality are known to affect mosquito reproductive success, but their combined effect has received little attention. Our study tested the hypothesis that the interaction between blood meal source and seasonality could affect mosquito fitness, a key factor that could play a role in seasonal host shift. By examining this interaction, we aimed to assess whether physiological processes like reproduction could be considered a candidate for shaping seasonal feeding patterns in Culex quinquefasciatus.

Added in the revised version: If seasonal host shift were driven by fitness advantages, we would expect mosquitoes to exhibit higher reproductive output when feeding on the predominant host in each season—that is, greater fecundity and fertility with bird blood in summer and with mammal blood in autumn. However, our findings did not fully align with these expectations. While mosquitoes fed on birds maintained high reproductive success in both summer and autumn, those fed on mammals exhibited a significant decline in reproductive output during autumn. This suggests that, rather than conferring a reproductive advantage in this season, mammalian blood may become a suboptimal resource under autumn conditions, potentially due to metabolic constraints or differences in nutrient assimilation.

Added in the revised version: The effect of blood meal on fitness observed in this study (greater in bird-fed mosquitoes) is consistent with previous findings on Cx. quinquefasciatus (Richards et al., 2012; Telang and Skinner, 2019) and other Culex (Demirci et al., 2014) and Aedes (Harrison et al., 2021) species. Avian blood is known to be nutritionally richer than mammalian blood, particularly in proteins and essential amino acids (e.g., isoleucine) required for egg development (Alto et al., 2014). Beyond intrinsic differences in blood composition, environmental factors such as temperature and photoperiod (seasonality) may further modulate how nutrients from a blood meal are allocated. While most of the nutrients obtained from a blood meal are directed toward egg production, a portion is also required for non-reproductive physiological functions (Rivera-Pérez et al., 2017), and the extent to which these resources are utilized may depend on external conditions such as temperature (Briegel and Timmermann, 2001).

Added in the revised version: Previous studies have shown that temperature regulates oogenesis and influences how mosquitoes allocate proteins obtained from a blood meal. In Aedes aegypti, for example, the proportion of blood-derived proteins allocated to egg production varies with temperature: at low temperatures, approximately 50% of these proteins are used for oogenesis, whereas at higher temperatures, this proportion drops to 25% (Briegel and Timmermann, 2001). If a similar mechanism occurs in Cx. quinquefasciatus, it could partially explain why mammal-fed mosquitoes exhibited a more substantial reproductive decline in autumn. Under lower temperatures, a greater proportion of proteins may be required for egg development, which could disproportionately impact individuals feeding on nutritionally poorer blood sources.

Added in the revised version: In addition to this blood meal-seasonal effect, we also found a significant interaction between blood meal source and gonotrophic cycle. This result suggests that the impact of blood sources is not uniform across cycles. If blood meal source were the only determinant of reproductive success, we would expect its effects to be consistent across cycles, with mosquitoes fed on birds maintaining higher reproductive output than those fed on mammals in a stable pattern. However, our findings showed that while bird-fed mosquitoes sustained similar reproductive output across cycles, those fed on mammals exhibited a significant decline in fertility in the second cycle. Since each cycle involves a new blood meal, this suggests that factors beyond blood nutritional quality contribute to the observed pattern, disproportionately affecting mosquitoes feeding on mammals in later cycles.

Added in the revised version: A decline in fertility across gonotrophic cycles has also been documented in other Cx. quinquefasciatus populations fed on both birds and mammals (Richards et al., 2012; Telang and Skinner, 2019), suggesting that reproductive output varies not only with seasonality but also with the number of successive cycles. One potential explanation for this decline is age-related reproductive constraints. With each successive cycle, ovarian follicular degeneration increases, leading to a reduction in the number of viable eggs per raft (Awahmukalah and Brooks, 1985). Additionally, given that several days separate the first and second cycles, physiological aging could contribute to this pattern (McCann et al., 2009). Interestingly, bird-fed mosquitoes maintained stable fertility across cycles, a pattern that has also been observed in previous studies (Bennett, 1970). While fertility often declines in later cycles, some studies have reported no change or even an increase in egg production under certain conditions (Bennett, 1970). This variability suggests that the effects of successive gonotrophic cycles on fertility may depend on multiple factors, including blood meal quality. Since avian blood is known to be nutritionally richer than mammalian blood, mosquitoes feeding on birds may have access to sufficient resources to sustain egg production across cycles, even as they age. However, further research is needed to determine whether specific nutrient components play a role in maintaining reproductive output over time. Similarly, body size has been shown to influence reproductive patterns in some mosquito species. In Aedes aegypti, for instance, larger females exhibit a rapid decline in fecundity after the first cycle, whereas smaller females maintain consistent egg production across early cycles before experiencing a significant drop (Briegel et al., 2002). While we did not assess body size in this study, these findings underscore the complexity of reproductive trade-offs in mosquitoes and the importance of considering multiple interacting factors when examining reproductive dynamics across cycles.

Added in the revised version: Together, these results indicate that both blood meal source and seasonality influence mosquito reproductive output. Moreover, hatchability did not differ among treatments, indicating that differences in egg viability did not drive these effects. This reinforces the idea that the main differences in reproductive output stem from factors affecting fecundity and fertility rather than post-oviposition processes.

Added in the revised version: While our findings confirm an interaction effect between blood meal source and seasonality, several caveats must be considered. First, our study was conducted under controlled laboratory conditions, which may not fully capture the environmental variability experienced by mosquitoes in the field. Factors such as fluctuating temperature, natural light cycles, and variable humidity levels could influence mosquito reproductive performance in ways not accounted for in our setup. However, as this is the first study of its kind, future research should aim to incorporate more ecologically realistic conditions. A second limitation commonly encountered in ecological and behavioural studies, is determining an appropriate sample size for detecting biologically meaningful effects (Taborsky, 2010). To address this, we repeated the assay across three replicates and conducted a post-hoc power analysis, which suggested that our sample size was adequate given the presence of interaction effects among explanatory variables. Finally, it remains possible that Cx. quinquefasciatus populations in Argentina do not exhibit the same host-switching behaviour observed in US populations, potentially influencing the relevance of our findings for different ecological contexts. However, previous studies have reported seasonal shifts in host selection and arbovirus transmission patterns in Argentina (Beranek, 2019; Spinsanti et al., 2008), suggesting that similar dynamics may be at play. More field-based studies will be essential to assess seasonal host variation in Cx. quinquefasciatus and its implications for arbovirus transmission.

Added in the revised version: In summary, our study investigates a critical aspect of mosquito reproductive success and its potential link to seasonal host selection in Culex mosquitoes, a group of significant epidemiological importance due to their role as bridge vectors between avian and human hosts. While we detected a significant interaction between blood meal source and seasonality, our findings do not support the idea that reproductive fitness alone explains seasonal host shifts. Instead, they suggest that the relationship between mosquito reproduction and host selection is more complex, likely influenced by multiple ecological and physiological factors beyond fitness constraints. Given this complexity, alternative explanations for seasonal host shifts should be considered. Genetic predisposition, midgut microbiota composition, physiological constraints, or other biological mechanisms unrelated to fitness may also contribute to shaping host-use patterns (Gervasi et al., 2016; Yan et al., 2018). Investigating these aspects will be crucial to fully understand the drivers of host selection in mosquitoes. Future research should aim to unravel the intricate relationship between feeding patterns, reproductive physiology, and environmental conditions. Refining our understanding of these dynamics will improve vector ecology models and enhance predictions of how seasonal variations in host selection may influence pathogen transmission.

Removed in the revised version: Our results show there is an interaction effect between blood meal source and seasonality on reproductive outputs of Cx. quinquefasciatus. While our model hypothesized that fecundity and fertility would be higher in bird-fed mosquitoes than in mammal-fed mosquitoes during summer, and the reverse would occur in autumn, our findings revealed the opposite pattern. Regardless the season, fecundity and fertility in Culex quinquefasciatus were higher in bird-fed mosquitoes compared to those fed on mammals. This pattern aligns with those findings from previous studies (Akoh et al., 1993; Richards et al., 2012; Telang & Skinner, 2019). This trend is not unique to Cx. quinquefasciatus but has also been seen in other Culex species (Shroyer & Silvery, 1972) and Aedes mosquitoes (Harrison et al., 2021). The commonly accepted explanation for this pattern suggests that the nutritional quality of avian blood, with its nucleated and larger erythrocytes, provides greater protein and energy, thereby enhancing these reproductive outputs (Alto et al., 2014). However, our findings offer a mixed perspective by highlighting the interaction between blood meal source and seasonality that further might influence fecundity and fertility.

Removed in the revised version: The significant interaction observed between blood meal source and seasonality revealed that mammal-fed mosquitoes had lower fecundity and fertility in autumn compared to summer, while bird-fed mosquitoes maintained similar reproductive outputs across both seasons. This finding diverges from the expectation of consistently higher fecundity and fertility in bird-fed mosquitoes and consistently lower fecundity and fertility in mammal-fed mosquitoes. This suggests that environmental factors such as temperature and photoperiod, which have been shown to affect mosquito fitness (Mordecai et al., 2019; Abouzied, 2017; Mogi, 1992), may interact with blood meal source in complex ways. Although the effects of temperature and photoperiod are generally coupled and follow a unimodal trend with an optimal peak at 28°C (Mordecai et al., 2019), our results indicate that the interaction between seasonality and blood source might play a critical role in determining reproductive success.

Removed in the revised version: Additionally, our study found a significant interaction between blood meal source and gonotrophic cycle on fertility. For mammal-fed mosquitoes, fertility was significantly lower in the second gonotrophic cycle compared to the first, a pattern commonly reported in other studies where the first and second cycles are generally the most productive (Awahmukalah & Brooks, 1985; Christiansen-Jucht et al., 2015). This declining fertility pattern has been documented in Cx. quinquefasciatus fed on both birds and mammals (Richards et al., 2012; Telang & Skinner, 2019). However, our results showed that this decline in fertility was not present in bird-fed mosquitoes, indicating a potential role of blood meal source in moderating reproductive output across gonotrophic cycles.

Removed in the revised version: We also investigated whether hatchability could account for the observed variations in fertility among different treatments. Our results indicate no difference in hatchability among treatments.

Removed in the revised version: While our findings corroborate an interaction effect between blood meal source and seasonality, they are not without caveats, as same as other studies. First, our research was conducted under controlled conditions in a laboratory incubator, which implies that our results may differ from those of field studies. Some of the main limitations of this approach include the use of artificial, constant light and temperature rather than a fluctuating environment, and the use of constant humidity, which could affect mosquito performance. However, this is the first study of its kind, and future investigations should aim for more realistic setups. A second limitation, commonly encountered in ecological and behavioral studies, is determining an appropriate sample size for each treatment, as this can impact the ability to detect desired effects or interactions (Taborsky, 2010). To address this issue, we conducted an a priori power analysis to determine the minimum required sample size and repeated the assay multiple times (three replicates). Furthermore, the presence of interaction effects among our explanatory variables suggests that our chosen sample size may have been adequate. Lastly, another aspect of this research that warrants further discussion is the underlying hypothesis that Culex mosquitoes change their feeding habits from birds to mammals. Given that most species are ornithophagic, why is this shift not directed from one bird host to another? In fact, both patterns—bird-to-mammal (Burkett-Cadena et al., 2011) and bird-to-bird (Kent et al., 2009) shifts—are mentioned in the literature. Future studies are aimed to clarify the complex feeding behaviour of Culex mosquitoes regarding seasonal switching, which may depend on several factors, such as habitat, as evidenced in U.S. Culex populations (Caranci, 2010).

Removed in the revised version: In summary, our study investigates the critical feature of host shifting patterns in Culex mosquitoes, which holds significant importance due to their role as bridge vectors between bird and human hosts. Although our study detects a significant interaction between blood meal source and seasonality, we observed the opposite expected trend, suggesting a complex interplay of factors affecting mosquito reproduction and its relationship with host use. We could consider three possible explanations for our results. Firstly, it is possible that additional factors, such as genetic mechanisms or variations in midgut microbiota composition play a role in influencing host use patterns. Secondly, host switching might be explained by factors related to vector biology other than fitness. Lastly, it remains a possibility that Argentinian Cx. quinquefasciatus populations do not exhibit the same host-switching behaviour observed in US populations, and the observed interaction effects may be linked to different mosquito behaviours. However, Spinsanti et al. (2008) found a seasonal variation in WNV human cases, and Beranek (2019) also observed that in autumn, the proportion of blood meals was higher than in summer for Culex quinquefasciatus. Therefore, more in-field studies are needed to assess seasonal host variation in Culex mosquitoes from Argentina.

Removed in the revised version: Further investigations are warranted to corroborate our findings and gain a deeper understanding of the interaction between mosquito feeding pattern and seasonality. Last, these investigations will provide valuable insights into mosquito-borne disease transmission and forecasting its emergence.

Removed in the revised version: Supporting information